# Homework 3: More String, Lists, and Python Objects (18 pts)

name: Annie Ye

email: yuhanye@umich.edu

This homework assignment took me 8 hours in total to complete. (Please help us to gauge the difficulty of the assignment.)

## Collaboration Disclosure

In the cell below, please list *everyone* with whom you discussed any of the homework problems, excluding only the GSIs and the course instructor. 

If you did not discuss the homework with anyone else, write __"I did not discuss this homework with anyone."__

Even if you discuss questions with other, the code you submit must be only yours. All work is checked with the [MOSS plagiarism detector](https://theory.stanford.edu/~aiken/moss/).

--- "I did not discuss this homework with anyone."

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. 


Before submitting, please make sure to __Cells->Run All__ executes without errors; errors in your code translate directly to point deductions. 
In general, you don't need to do explicitly raise errors (e.g. with the ```raise``` function) if we don't ask you to in the problem statement.
However, even in cases where we ask you to check for errors, your submission should not contain any examples of your functions actually raising those errors.

Note that many parts of this homework where you are expected to type in code will have ```NotImplementedError()``` as a placeholder. You need to delete this function and replace it with your own code.

## Importing libaries

Unless specifically prohibited, you can import any Python library that comes included with Python (e.g. ```math```, ```datetime```) or that has been used in the lecture or homework prompts. You do not need any other libraries to complete the homework. Note that the graders will not install any additional libraries to grade an assignment.


<!-- ## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking: -->

## Homework tips 

1. **Start early!** If you run into trouble installing things or importing packages, it’s
best to find those problems well in advance, not the night before your assignment is
due when we cannot help you!

2. **Make sure you back up your work!** At a minimum, do your work in a Dropbox
folder. Better yet, use git, which is well worth your time and effort to learn.

3. **Be careful to follow directions!** Remember that Python is case sensitive. If
you are ask you to define a function called my_function and you define a function
called My_Function, you will not receive full credit. You may want to copy-paste
the function names below to make sure that the functions in your notebook match.

## Error checking

You do not need to do error checking (raising errors, etc.) in your code unless we explicitly ask you to so in a problem.

## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking:


## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. You will notice some `read-only` cells in the assignment that contain `assert` statements. These are tests that your code must pass for your solution to be correct. If any of the tests fail, you will get an python error and not get points for that question. 

**Note:** The tests shown not are not comprehensive; additional tests will be used at grading time. You are encouraged to read the problem carefully and verify your code covers all possible cases.

**Be careful:** If a jupyter notebook cell takes longer than `60s` to run, the autograder will not grade it and you will receive zero credit for that question.

```python
import collections
import string
import itertools
```

---
## Question 1: Infinite sequences (8 pts)
For each of the problems below, write a function which generates the given infinite sequence. We should be able to use your generators to access any entry of the sequence no matter how deep.

**1(a)** (2 pts) The prime numbers are

$$2, 3, 5, 7, 11, 13, 17, \dots$$

Give a generator for the primes. (Note: many algorithms exist for this problem. Yours should be efficient enough that we can use it to generate reasonably large prime numbers.)

<!-- grade_id: primes, points: 0, grade: false, solution: true -->
```python
def primes():
    # YOUR CODE HERE
    primes = [2]
    k = 3
    yield 2
    while k > 2:
        for prime in primes:
            if k % prime == 0:
                break
        else:
            primes.append(k)
            yield k
        k += 1        
```

<!-- grade_id: correct_prime, points: 2, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
p = primes()
plist = [2,3,5,7,11,13,17,19,23,29]
for i in range(10):
    assert next(p) == plist[i]
```

**1(b)** (3 pts) The *ruler sequence* is

$$1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7\dots$$

The first number (1) appears once; the next two numbers (2 and 3) appear twice, the next three numbers appear three times, etc.

Hint: You are highly encouraged to use relevant functions from itertools to make your solution simple

<!-- grade_id: ruler, points: 0, grade: false, solution: true -->
```python
import itertools

def ruler():
    # YOUR CODE HERE
    k = 1
    num = 1
    count = 1
    while True:
        for i in range(num):
            yield from itertools.repeat(k, count)
            k += 1
            i += 1
        count += 1
        num += 1
```

<!-- grade_id: correct_ruler, points: 3, grade: true, solution: false -->
```python
r = ruler()
ruler_numbers = [1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7, 8, 8, 8, 8]
for n in ruler_numbers:
    assert next(r) == n
```

**1(c)** (3 pts) The look-and-say sequence  

$$1, 11, 21, 1211, 111221, 312211, 13112221, 1113213211, \dots$$

is generated as follows:

1. The first entry is one.
2. The second entry is generated by reading the first entry aloud: "one one"
3. The third entry is generated by reading the second entry aloud: "two ones"
4. The fourth entry is generated by reading the third entry aloud: "one two and one one"
5. The fifth entry is therefore "one one, one two, and two ones"

... and so forth. Note that each entry of the sequence should be a `str` object.

<!-- grade_id: look_say, points: 0, grade: false, solution: true -->
```python
def look_say():
    # YOUR CODE HERE
    count = 1
    str1 = '1'
    yield str1
    while True:
        str2 = ''
        count = 1
        for k in range(1, len(str1)):
            if str1[k] == str1[k - 1]:
                count += 1
            else:
                str2 += str(count) + str1[k - 1]
                count = 1
        str2 += str(count) + str1[- 1]        
        yield str2
        str1 = str2            
```

```python
import time

def validate_look_say(prev, current):
    """Check if `current` is a valid description of `prev` in the look-and-say sequence."""
    description = ''
    count = 1
    for i in range(1, len(prev)):
        if prev[i] == prev[i - 1]:
            count += 1
        else:
            description += str(count) + prev[i - 1]
            count = 1
    description += str(count) + prev[-1]  # Add last group

    # Print whether the generated term matches the expected term
    if description == current:
        print("Validation successful: Terms match!")
    else:
        print("Validation failed: Terms do not match!")

    return description == current

def test_look_say_sequence_with_timing(limit=100):
    """Test the look-and-say sequence up to `limit` terms with timing."""
    s = look_say()
    prev_term = next(s)
    print(f"First term: {prev_term}")  # Debug: Print the first term
    
    start_time = time.time()  # Start timing

    for i in range(1, limit):
        current_term = next(s)

        # Validate the current term based on the previous one
        assert validate_look_say(prev_term, current_term), f"Test failed at term {i+1}"
        
        # Print every term to track progress and time
        if i % 10 == 0:  # Adjust this to print less frequently if needed
            elapsed_time = time.time() - start_time  # Calculate time elapsed
            print(f"Validated {i} terms. Time elapsed: {elapsed_time:.2f} seconds.")

        prev_term = current_term
    
    print(f"Test passed: Successfully generated and validated {limit} terms in {time.time() - start_time:.2f} seconds!")

# Run the test for 100 terms
test_look_say_sequence_with_timing(100)
```

Output:
```
First term: 1
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validated 10 terms. Time elapsed: 0.00 seconds.
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validated 20 terms. Time elapsed: 0.00 seconds.
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Terms match!
Validation successful: Ter
```

<!-- grade_id: correct_look_say, points: 3, grade: true, solution: false -->
```python
s = look_say()
look_say_first_ten = ['1',
 '11',
 '21',
 '1211',
 '111221',
 '312211',
 '13112221',
 '1113213211',
 '31131211131221',
 '13211311123113112211']
for i in range(10):
    s_i = next(s)
    assert isinstance(s_i, str)
    assert s_i == look_say_first_ten[i]
```

### Question 2: Simple ciphers (5 pts)

A *cipher* is an algorithm for encrypting or decrypting a text message, called *plaintext*, into *ciphertext*. In this exercise, we will assume that all plaintext to be encrypted or decrypted are composed of the lower-case Roman alphabet, without any punctuation or whitespace. Examples of such messages could be

```attackatdawn```

or 

```iamajellyfilleddonut```.

<!-- grade_id: enc_caesar, points: 0, grade: false, solution: true -->
```python
def enc_caesar(s: str, k: int) -> str:
    '''
    Encrypt the message s using Ceasar cipher with key k.
    
    >>> enc_caeasar('thequickbrownfoxjumpsoverthelazydog', 1)
    'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'
    '''
    # YOUR CODE HERE
    enc_str = ''
    if k % 26 == 0:
        return s
    
    k = k % 26
    
    for i in range(len(s)):
        new_char = chr(((ord(s[i]) - ord('a') + k) % 26) + ord('a'))
        enc_str += new_char
    return enc_str
```

<!-- grade_id: correct_enc_caesar, points: 2, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests:
    assert enc_caesar(msg, k) == enc
```

```python
enc_tests = [
    ('thequickbrownfoxjumpsoverthelazydog', 0, 'thequickbrownfoxjumpsoverthelazydog'),
    ('packmyboxwithfivedozenliquorjugs', 5, 'ufhprdgtcbnymknajitejsqnvztwozlx'),
    ('abcdefghijklmnopqrstuvwxyz', 13, 'nopqrstuvwxyzabcdefghijklm'),
    ('abcdefghijklmnopqrstuvwxyz', 25, 'zabcdefghijklmnopqrstuvwxy'),
    ('thequickbrownfoxjumpsoverthelazydog', 26, 'thequickbrownfoxjumpsoverthelazydog'),
    ('packmyboxwithfivedozenliquorjugs', 52, 'packmyboxwithfivedozenliquorjugs'),
    ('packmyboxwithfivedozenliquorjugs', 27, 'qbdlnzcpyxjuigjwfepafomjrvpskvht'),
    ('abcdefghijklmnopqrstuvwxyz', 100, 'wxyzabcdefghijklmnopqrstuv'),
    ('thequickbrownfoxjumpsoverthelazydog', -1, 'sgdpthbjaqnvmenwitlornudqsgdkzyxcnf'),
    ('thequickbrownfoxjumpsoverthelazydog', -25, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'),
    ('packmyboxwithfivedozenliquorjugs', -27, 'ozbjlxanwvhsgehudcnydmkhptnqitfr'),
    ('thequickbrownfoxjumpsoverthelazydog', -100, 'xliuymgofvsarjsbnyqtwszivxlipedchsk'),
]

for msg, k, expected_output in combined_tests:
    output = enc_caesar(msg, k)
    assert output == expected_output, f"Test failed for input: '{msg}' with key {k}. Got {output}, expected {expected_output}"

print("All enc_tests cases passed!")
```

<!-- grade_id: dec_caesar, points: 0, grade: false, solution: true -->
```python
def dec_caesar(s: str, k: int) -> str:
    '''
    Decrypt the message s using Ceasar cypher with key k.
    
    >>> dec_caeasar('uifrvjdlcspxogpykvnqtpwfsuifmbazeph', 1)
    'thequickbrownfoxjumpsoverthelazydog'
    '''
    # YOUR CODE HERE
    dec_str = ''
    if k % 26 == 0:
        return s
    
    k = -(k % 26)
    
    for i in range(len(s)):
        dec_str += chr(((ord(s[i]) - ord('a') + k) % 26) + ord('a'))
    return dec_str
```

<!-- grade_id: correct_dec_caesar, points: 3, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests: 
    assert dec_caesar(enc, k) == msg
```

```python
dec_tests = [
    ('uifrvjdlcspxogpykvnqtpwfsvifmbazep', 1, 'thequickbrownfoxjumpsoveruhelazydo'),  # Shifted by 1
    ('ufhpdbctbvnkyonajitqjsvqnytzpynx', 5, 'packywxowqiftjivedolenqlitouktis'),  # Shifted by 5
    ('nopqrstuvwxyzabcdefghijklm', 13, 'abcdefghijklmnopqrstuvwxyz'),  # Shifted by 13 (ROT13)
    ('zabcdefghijklmnopqrstuvwxy', 25, 'abcdefghijklmnopqrstuvwxyz'),  # Shifted by 25
    ('thequickbrownfoxjumpsoverthelazydog', 26, 'thequickbrownfoxjumpsoverthelazydog'),  # No shift
    ('packmyboxwithfivedozenliquorjugs', 52, 'packmyboxwithfivedozenliquorjugs'),  # Shifted by 52 (no change)
    ('qbdlnzcpyxjvkgjwfpadenmjrvsktvht', 27, 'packmyboxwiujfiveozcdmliqurjsugs'),  # Shifted by 27 (equivalent to 1)
    ('mnopqrstuvwxyzabcdefghijkl', 100, 'qrstuvwxyzabcdefghijklmnop'),  # Shifted by 100 (100 % 26 = 22)
    
    ('sgdpthbjaqnvemwitlornudqsgkyxcnf', -1, 'thequickbrowfnxjumpsoverthlzydog'),  # Negative shift by -1
    ('sgdpthbjaqnvemwitlornudqsgkyxcnf', -25, 'rfcosgaizpmudlvhsknqmtcprfjxwbme'),  # Equivalent to shifting by 1
    ('sgdpthbjaqnvemwitlornudqsgkyxcnf', -26, 'sgdpthbjaqnvemwitlornudqsgkyxcnf'),
    ('ozbllxanvhsgeuhndcygkhptkhtniftr', -27, 'pacmmybowithfvioedzhliquliuojgus'),  # Shifted by -1
    ('pcfoqwezdjmtlupfnslrkqtpdyjxydmn', -100, 'lybkmsavzfiphqlbjohngmplzuftuzij'),  # Large negative shift (-100 % 26 = -22)
]

# Run the decryption test cases
for msg, k, expected_output in dec_tests:
    output = dec_caesar(msg, k)
    assert output == expected_output, f"Test failed for decryption of '{msg}' with key {k}. Got {output}, expected {expected_output}"

print("All dec_tests cases passed!")
```

### Question 3 (5 pts) Valid String.
Given a string s containing just the characters '(', ')', '{', '}', '[' and ']', determine if the input string is valid.

An input string is valid if:

1) Open brackets must be closed by the same type of brackets.
2) Open brackets must be closed in the correct order.
3) Every close bracket has a corresponding open bracket of the same type.


<!-- grade_id: is_valid, points: 0, grade: false, solution: true -->
```python
def is_valid(s):
    bracket_map = {')': '(', '}': '{', ']': '['}
    stack = []

    for char in s:
        if char in bracket_map.values():
            stack.append(char)
        elif char in bracket_map:
            if not stack or stack[-1] != bracket_map[char]:
                return False
            stack.pop()
        else:
            return False

    return not stack
```

<!-- grade_id: correct_is_valid, points: 5, grade: true, solution: false -->
```python
assert is_valid('()') == True
assert is_valid('([{]}') == False
assert is_valid('([{}])') == True
```

```python
assert is_valid('[)(]') == False
assert is_valid('[[]]') == True
assert is_valid("()[]{}") == True
assert is_valid("({[]})") == True
assert is_valid("((((((((()))))))))") == True
assert is_valid("()(){}[]") == True
assert is_valid("{[()()]}") == True
assert is_valid("") == True
assert is_valid("(") == False
assert is_valid(")") == False
assert is_valid("(]") == False
assert is_valid("([)]") == False
assert is_valid("{[}") == False
assert is_valid("{[(()]") == False
assert is_valid("(((()))") == False
assert is_valid("(((([]))))") == True
assert is_valid("([{}((){})])") == True
assert is_valid("{([][)]}") == False
```

```python

```

```python

```